Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tianxiao Hu"
COLLABORATORS = ""
```

---

# Homework 4: Spam/Ham Classification
## Feature Engineering, Logistic Regression, Cross Validation
## Due Date: 11/1/18, 11:59PM
## Course Policies

Here are some important course policies. These are also located at
http://www.ds100.org/fa18/.

**Collaboration Policy**

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your solution.

## This Assignment
In this homework, you will use what you've learned in class to create a classifier that can distinguish spam (junk or commercial or bulk) emails from ham (non-spam) emails. In addition to providing some skeleton code to fill in, we will evaluate your work based on your model's accuracy and your written responses in this notebook.

After this homework, you should feel comfortable with the following:

- Feature engineering with text data
- Using sklearn libraries to process data and fit models
- Validating the performance of your model and minimize overfitting
- Generating and analyzing precision recall curves

## Warning
We've tried our best to filter the data for anything blatantly offensive as best as we can, but unfortunately there may still be some examples you may find in poor taste. If you encounter these examples and believe it is inappropriate for students, please let a TA know and we will try to remove it for future semesters. Thanks for your understanding!

## Regarding Submissions - IMPORTANT, PLEASE READ


For this assignment and future assignments (homework and projects) you will also submit your free response and plotting questions to Gradescope. To do this, you can download as PDF; there are two ways to do this:

 

1. File > Print Preview ----> Save website as PDF
2. Control/Command + P ----> Save website as PDF
 

If you are having trouble with text being cut off in the generated PDF:

- For cells containing text surrounded in triple quotes (e.g. """ Hello world """), you can press enter in the middle of the string to push the text to a new line so that all the text stays within the box.


You are responsible for submitting and tagging your answers in Gradescope. For each free response and plotting question, please include:

1. Relevant code used to generate the plot or inform your insights
2. The written free response or plot

# Part I - Initial Analysis

```python
import numpy as np
import pandas as pd

import matplotlib.pyplot as plt
%matplotlib inline

import seaborn as sns
sns.set(style = "whitegrid", 
        color_codes = True,
        font_scale = 1.5)
```

### Loading in the Data

The dataset consists of email messages and their labels (0 for ham, 1 for spam). 

Your labeled dataset contains 8348 labeled examples, and the test set contains 1000 unlabeled examples.

Run the following cells to load in the data into DataFrames.

The `train` DataFrame contains labeled data that you will use to train your model. It contains four columns:

1. `id`: An identifier for the training example.
1. `subject`: The subject of the email
1. `email`: The text of the email.
1. `spam`: 1 if the email was spam, 0 if the email was ham (not spam).

The `test` DataFrame contains another set of 1000 unlabeled examples. You will predict labels for these examples and submit your predictions to Kaggle for evaluation.

```python
from utils import fetch_and_cache_gdrive
fetch_and_cache_gdrive('1SCASpLZFKCp2zek-toR3xeKX3DZnBSyp', 'train.csv')
fetch_and_cache_gdrive('1ZDFo9OTF96B5GP2Nzn8P8-AL7CTQXmC0', 'test.csv')

original_training_data = pd.read_csv('data/train.csv')
test = pd.read_csv('data/test.csv')

# Convert the emails to lower case as a first step to processing the text
original_training_data['email'] = original_training_data['email'].str.lower()
test['email'] = test['email'].str.lower()

original_training_data.head()
```

# Question 1a
First let's check if our data contains any nan values. *Fill in the cell below to print whether any of the columns contain nan values.* If there are nan values, replace them with the appropriate filler values. In other words, a nan value in the subject column should be replaced with an empty string. 

Note that while there are no nan values in the spam column, we should be careful when replacing nan values when they are the labels. Doing so without consideration may introduce significant bias into our model when fitting.

<!-- grade_id: cell-b1fb39d9b651ca1b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
for column in original_training_data.columns:
    print(column, original_training_data[column].hasnans)
    
original_training_data.fillna(value="", inplace=True)
test.fillna(value="", inplace=True)

print(original_training_data["subject"].hasnans)
```

# Question 1b
In the cell below, print the text of the first ham and the first spam email in the original training set. Then, discuss one thing you notice that is different between the two that might relate to the identification of spam.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Print the text of the first ham and the first spam emails. Then, fill in your response in the q01 variable:
first_ham = original_training_data[original_training_data["spam"] == 0].iloc[0].email
first_spam = original_training_data[original_training_data["spam"] == 1].iloc[0].email

# YOUR CODE HERE
print(first_ham)
print(first_spam)
# raise NotImplementedError()
```

<!-- grade_id: q1-tests, points: 1, grade: true, solution: false -->
```python
# This is a cell with just a comment but don't delete me if you want to get credit.
```

<!-- grade_id: q1-written, points: 2, grade: true, solution: true -->
The spam email is in HTML format while the ham email is in text format.

## Training Validation Split
The training data we downloaded is all the data we have available for both training models and **validating** the models that we train.  We therefore need to split the training data into separate training and validation datsets.  You will need this **validation data** to validate your model once you are finished training. Note that we set the seed (random_state) to 42. This will produce a pseudo-random sequence of random numbers. Do not modify this in the following questions, as our assert statements depend on this random seed.

```python
from sklearn.model_selection import train_test_split

[train, val] = train_test_split(original_training_data, test_size=0.1, random_state=42)
```

# Basic Feature Engineering

We would like to take the text of an email and predict whether the text is ham or spam. This is a *classification* problem, so we can use logistic regression to make a classifier. Recall that to train an logistic regression model we need a numeric feature matrix $\Phi$ (pronounced phi as in wifi) and corresponding binary labels $Y$.  Unfortunately, our data are text, not numbers. To address this, we can create numeric features derived from the email text and use those features for logistic regression.

Each row of $\Phi$ is derived from one email example. Each column of $\Phi$  is one feature. We'll guide you through creating a simple feature, and you'll create more interesting ones when you are trying to increase your accuracy.

# Question 2

Create a function called `words_in_texts` that takes in a list of `words` and a pandas Series of email `texts`. It should output a 2-dimensional NumPy array containing one row for each email text. The row should contain either a 0 or a 1 for each word in the list: 0 if the word doesn't appear in the text and 1 if the word does. For example:

```python
>>> words_in_texts(['hello', 'bye', 'world'], 
                   pd.Series(['hello', 'hello world hello']))

array([[1, 0, 0],
       [1, 0, 1]])
```

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def words_in_texts(words, texts):
    '''
    Args:
        words (list-like): words to find
        texts (Series): strings to search in
    
    Returns:
        NumPy array of 0s and 1s with shape (n, p) where n is the
        number of texts and p is the number of words.
    '''
    indicator_array = np.zeros((len(texts), len(words)))
    for i in range(len(texts)):
        for j in range(len(words)):
            if(words[j] in texts.iloc[i]):
                indicator_array[i][j] = 1
    # YOUR CODE HERE
#     raise NotImplementedError()
    return indicator_array
```

<!-- grade_id: q2-tests, points: 2, grade: true, solution: false -->
```python
# If this doesn't error, your function outputs the correct output for this example
assert np.allclose(words_in_texts(['hello', 'bye', 'world'], 
                                  pd.Series(['hello', 'hello world hello'])),
                   np.array([[1, 0, 0], 
                             [1, 0, 1]]))

assert np.allclose(words_in_texts(['a', 'b', 'c', 'd', 'e', 'f', 'g'], 
                                  pd.Series(['a b c d e f g', 'a', 'b', 'c', 'd e f g', 'h', 'a h'])),
                   np.array([[1,1,1,1,1,1,1], 
                             [1,0,0,0,0,0,0],
                             [0,1,0,0,0,0,0],
                             [0,0,1,0,0,0,0],
                             [0,0,0,1,1,1,1],
                             [0,0,0,0,0,0,0],
                             [1,0,0,0,0,0,0]]))
```

# Basic EDA

Now we need to identify some features that allow us to tell spam and ham emails apart. One idea is to compare the distribution of a single feature in spam emails to the distribution of the same feature in ham emails. If the feature is itself a binary indicator, such as whether a certain word occurs in the text, this amounts to comparing the proportion of spam emails with the word to the proportion of ham emails with the word.


# Question 3a

Create a bar chart comparing the proportion of spam and ham emails containing certain words. It should look like the following plot (which was created using `sns.barplot`), but you should choose your own words as candidate features. Make sure to use the training set (after splitting).

![training conditional proportions](./images/training_conditional_proportions.png "Class Conditional Proportions")

Hint:
- You can use DataFrame's `.melt` method to "unpivot" a DataFrame. See the following code cell for example

<!-- grade_id: q3a-answer, points: 2, grade: true, solution: true -->
```python
train=train.reset_index() # We must do this in order to preserve the ordering of emails to labels for words_in_texts
spam_words = ["html", "head", "body", "please", "offer", "money", "sale"]
spam_array = words_in_texts(spam_words, train["email"])
df = pd.DataFrame(spam_array, columns=spam_words)
df.insert(loc=0, column="type", value=train["spam"].apply(lambda x: "spam" if x == 1 else "ham"))
df = df.melt("type")
sns.barplot(x="variable", y="value", hue="type", data=df, ci=None)
plt.xlabel("Words")
plt.ylabel("Proportion of Emails")
plt.title("Frequency of Words in Spam/Ham emails")
plt.ylim(0, 1)
# YOUR CODE HERE
# raise NotImplementedError()
```

# Question 3b

When the feature is binary, it makes sense (as in the previous question) to compare the proportion of 1s in the two classes of email. Otherwise, if the feature can take on many values, it makes sense to compare the distribution under spam to the distribution under ham. Create a *class conditional density plot* like the one below (which was created using `sns.distplot`), comparing the distribution of a feature among all spam emails to the distribution of the same feature among all ham emails. You should use your training set (after splitting). **You may use the length of the email body or create your own feature.** If using length of the email body, please set the xlim to 50000.

![training conditional densities](./images/training_conditional_densities2.png "Class Conditional Densities")

<!-- grade_id: q3b-answer, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
ham_len = train[train["spam"] == 0]["email"].apply(len)
spam_len = train[train["spam"] == 1]["email"].apply(len)
sns.distplot(ham_len, hist=False)
sns.distplot(spam_len, hist=False)
plt.xlim(0, 50000)
plt.xlabel("Length of email body")
plt.ylabel("Distribution")
```

# Basic Classification

Notice that the output of `words_in_texts(words, train['email'])` is a numeric matrix containing features for each email. This means we can use it directly to train a classifier!

# Question 4

We've given you 5 words that might be useful as features to distinguish spam/ham emails. Use these words as well as the `train` DataFrame to create two NumPy arrays: `Phi_train` and `Y_train`.

`Phi_train` should be a matrix of 0s and 1s created by using your `words_in_texts` function on all the emails in the training set.

`Y_train` should be a vector of the correct labels for each email in the training set.

<!-- grade_id: q4-answer, points: 0, grade: false, solution: true -->
```python
some_words = ['drug', 'bank', 'prescription', 'memo', 'private']

Phi_train = words_in_texts(some_words, train["email"])
Y_train = train["spam"].values

# YOUR CODE HERE
# raise NotImplementedError()

Phi_train[:5], Y_train[:5]
```

<!-- grade_id: q4-tests, points: 2, grade: true, solution: false -->
```python
assert np.all(np.unique(Phi_train) == np.array([0, 1]))
assert np.all(np.unique(Y_train) == np.array([0, 1]))
assert Phi_train.shape[0] == Y_train.shape[0]
assert Phi_train.shape[1] == len(some_words)
```

# Question 5

Now we have matrices we can give to scikit-learn! Using the [`LogisticRegression`](http://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html) classifier, train a logistic regression model using `Phi_train` and `Y_train`. Then, output the accuracy of the model (on the training data) in the cell below. You should get an accuracy of around 0.75.

<!-- grade_id: q5-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LogisticRegression
model = LogisticRegression()
model = model.fit(Phi_train, Y_train)
training_accuracy = model.score(Phi_train, Y_train)

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q5-tests, points: 2, grade: true, solution: false -->
```python
assert training_accuracy > 0.72
```

# Question 6

That doesn't seem too shabby! But the classifier you made above isn't as good as this might lead us to believe. First, we are validating on the training set, which may lead to a misleading accuracy measure, especially if we used the training set to identify discriminative features. In future parts of this analysis, it will be safer to hold out some of our data for model validation and comparison.

Presumably, our classifier will be used for **filtering**, i.e. preventing messages labeled `spam` from reaching someone's inbox. Since we are trying  There are two kinds of errors we can make:
- False positive (FP): a ham email gets flagged as spam and filtered out of the inbox.
- False negative (FN): a spam email gets mislabeled as ham and ends up in the inbox.

These definitions depend both on the true labels and the predicted labels. False positives and false negatives may be of differing importance, leading us to consider more ways of evaluating a classifier, in addition to overall accuracy:

**Precision** measures the proportion $\frac{\text{TP}}{\text{TP} + \text{FP}}$ of emails flagged as spam that are actually spam.

**Recall** measures the proportion $\frac{\text{TP}}{\text{TP} + \text{FN}}$ of spam emails that were correctly flagged as spam. 

**False-alarm rate** measures the proportion $\frac{\text{FP}}{\text{FP} + \text{TN}}$ of ham emails that were incorrectly flagged as spam. 

The following image might help:

<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/2/26/Precisionrecall.svg/700px-Precisionrecall.svg.png" width="500px">

Note that a true positive (TP) is a spam email that is classified as spam, and a true negative (TN) is a ham email that is classified as ham. Answer the following questions in the cells below:

- (a) Suppose we have a classifier that just predicts 0 (ham) for every email. How many false positives are there? How many false negatives are there? Provide specific numbers using the training data from Question 4.
- (b) Suppose we have a classifier that just predicts 0 (ham) for every email. What is its accuracy on the training set? What is its recall on the training set?
- (c) What are the precision, recall, and false-alarm rate of the logistic regression classifier in Question 5? Are there more false positives or false negatives? 
- (d) Our logistic regression classifier got 75.6% prediction accuracy (number of correct predictions / total). How does this compare with predicting 0 for every email?
- (e) Given the word features we gave you above, name one reason this classifier is performing poorly.
- (f) Which of these two classifiers would you prefer for a spam filter and why? (N.B. there is no "right answer" here but be thoughtful in your reasoning).

<!-- grade_id: q6a-answer, points: 0, grade: false, solution: true -->
```python
# provide number of FP and FN, respectively,
# for a classifier that always predicts 0 (never predicts positive...)
zero_predictor_fp = 0
zero_predictor_fn = 1918

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q6a-tests, points: 2, grade: true, solution: false -->
```python
# This is a cell with just a comment but don't delete me if you want to get credit.
```

<!-- grade_id: q6b-answer, points: 0, grade: false, solution: true -->
```python
# provide training accuracy & recall, respectively,
# for a classifier that always predicts 0
zero_predictor_acc =  float(sum(Y_train == 0)) / len(Y_train)
zero_predictor_recall = 0
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q6b-tests, points: 2, grade: true, solution: false -->
```python
# This is a cell with just a comment but don't delete me if you want to get credit.
```

<!-- grade_id: q6c-answer, points: 0, grade: false, solution: true -->
```python
# provide training accuracy & recall, respectively,
# for logistic regression classifier from question 5
TP = float(sum(model.predict(Phi_train[Y_train == 1]) == 1))

logistic_predictor_precision = TP / sum(model.predict(Phi_train) == 1)
logistic_predictor_recall = TP / sum(Y_train == 1)
logistic_predictor_far = float(sum(model.predict(Phi_train[Y_train == 0]) == 1)) / sum(Y_train == 0)

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q6c-tests, points: 2, grade: true, solution: false -->
```python
# This is a cell with just a comment but don't delete me if you want to get credit.
```

<!-- grade_id: q6-written, points: 3, grade: true, solution: true -->
- (a) 0 false positives and 1918 false negatives
- (b) accuracy is 74.5% and recall is 0
- (c) precision is 64.2%, recall is 11.4% and false-alarm rate is 2.2%
- (d) just a little higher than a zero predictor
- (e) the feature can't predict whether the email is a junk or ham well
- (f) zero predictor for losing important emails can be awful

# Part II - Moving Forward

With this in mind, it is now your task to make the spam filter more accurate. In order to get full credit on the accuracy part of this assignment, you must get at least **88%** accuracy on the test set. To see your accuracy on the test set, you will use your classifier to predict every email in the `test` DataFrame and upload your predictions to Kaggle.

To prevent you from overfitting to the test set, you may only upload predictions to Kaggle twice per day. This means you should start early and rely on your **validation data** to estimate your Kaggle scores.  

Here are some ideas for improving your model:

1. Finding better features based on the email text. Some example features are:
    1. Number of characters in the subject / body
    1. Number of words in the subject / body
    1. Use of punctuation (e.g., how many '!' were there?)
    1. Number / percentage of capital letters 
    1. Whether the email is a reply to an earlier email or a forwarded email
1. Finding better words to use as features. Which words are the best at distinguishing emails? This requires digging into the email text itself. 
1. Better data processing. For example, many emails contain HTML as well as text. You can consider extracting out the text from the HTML to help you find better words. Or, you can match HTML tags themselves, or even some combination of the two.
1. Model selection. You can adjust parameters of your model (e.g. the regularization parameter) to achieve higher accuracy. Recall that you should use cross-validation to do feature and model selection properly! Otherwise, you will likely overfit to your training data.

You may use whatever method you prefer in order to create features. However, **you are only allowed to train logistic regression models and their regularized forms**. This means no random forest, k-nearest-neighbors, neural nets, etc.

We will not give you a code skeleton to do this, so feel free to create as many cells as you need in order to tackle this task. However, answering questions 7, 8, and 9 should help guide you.

---

**Note:** *You should use the **validation data** to evaluate your model and get a better sense of how it will perform on the Kaggle evaluation.*

---

# Question 7 (Feature/Model Selection Process)

In this following cell, describe the process of improving your model. You should use at least 2-3 sentences each to address the follow questions:

1. How did you find better features for your model?
2. What did you try that worked / didn't work?
3. What was surprising in your search for good features?

<!-- grade_id: q7-written, points: 6, grade: true, solution: true -->
1. I printed several ham emails and spam emails, read them and tried to find some difference between these two kinds of emails. I added some features according to tips above like the number of exclaimations in subject, number of words in the subject/body. Then I found the accuracy improved each time when a new feature is added (about 0.5% - 1% per feature).
2. I brainstormed and tried several features. I found the indicator of whether the email is a reply to former email(e.g. "Re:" is in the subject) contributes a lot for improvement of accuracy. However, whether the email is a forwarded email doesn't improve the accuracy significantly because there are very few samples that are forwarded email.
3. After I got in a mess with new features and still can't reach the required accuracy, I tried to combine former `word indicators` for prediction and it greatly improved the accuracy(5%+).

# Question 8 (EDA)

In the two cells below, show a visualization that you used to select features for your model. Include both

1. A plot showing something meaningful about the data that helped you during feature / model selection.
2. 2-3 sentences describing what you plotted and what its implications are for your features.

Feel to create as many plots as you want in your process of feature selection, but select one for the cells below.

**You should not show us a visualization just like in question 3.** Specifically, don't show us a bar chart of proportions, or a one-dimensional class conditional density plot. Any other plot is acceptable, as long as it comes with thoughtful commentary. Here are some ideas:

1. Consider the correlation between multiple features (look up correlation plots and `sns.heatmap`). 
1. Try to show redundancy in a group of features (e.g. `body` and `html` might co-occur relatively frequently, or you might be able to design a feature that captures all html tags and compare it to these). 
1. Use a word-cloud or another visualization tool to characterize the most common spam words.
1. Visually depict whether spam emails tend to be wordier (in some sense) than ham emails.

```python
import re

def feature_subject_upper(df):
    return df["subject"].apply(lambda x: sum([c.isupper() for c in x]))

def feature_subject_length(df):
    return df["subject"].apply(len)

def feature_subject_words(df):
    return df["subject"].apply(lambda x: len(re.findall(r'\w+', x)))

def feature_subject_exclaimation(df):
    return df["subject"].str.count("!")

def feature_re(df):
    return df["subject"].apply(lambda x: 1.0 if "Re:" in x else 0.0)

def feature_context_length(df):
    return df["email"].str.len()

def feature_context_words(df):
    return df["email"].apply(lambda x: len(re.findall(r'\w+', x)))

def feature_html_tag(df):
    return df["email"].str.count("/>")

day = ['mon', 'tue', 'wed', 'thur', 'fri', 'sat', 'sun', 'monday', 'tuesday', 'wednesday', 'thursday', 'friday', 'saturday', 'sunday']
month = ['january', 'february', 'march', 'april', 'may', 'june', 'july', 'august', 'september', 'october', 'november', 'december']

def feature_has_day(df, day):
    return df["email"].apply(lambda x: sum([d in x for d in day]))

def feature_has_month(df, month):
    return df["email"].apply(lambda x: sum([m in x for m in month]))
```

```python
context_length = feature_context_length(original_training_data)
context_words = feature_context_words(original_training_data)
```

```python
data_df = pd.DataFrame(columns=["context_length", "context_words"])
```

```python
data_df["context_length"] = context_length
data_df["context_words"] = context_words
```

```python
sns.jointplot(x="context_length", y="context_words", data=data_df, kind="scatter", xlim=(0, 12000), ylim=(0, 2000))
```

<!-- grade_id: q8-eda, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q8-commentary, points: 3, grade: true, solution: true -->
I draw a jointplot of the `number of words in email context` and `length of email context`. From the graph we can tell there is a redundancy between the two feature. They have positive correlation.

# Question 9 (Making a Precision-Recall Curve)

We can trade off between precision and recall. In most cases we won't be able to get both perfect precision (i.e. no false positives) and recall (i.e. no false negatives), so we have to compromise. For example, in the case of cancer screenings, false negatives are comparatively worse than false positives — a false negative means that a patient might not discover a disease until it's too late to treat, while a false positive means that a patient will probably have to take another screening.

Recall that logistic regression calculates the probability that an example belongs to a certain class. Then, to classify an example we say that an email is spam if our classifier gives it $\ge 0.5$ probability of being spam. However, *we can adjust that cutoff*: we can say that an email is spam only if our classifier gives it $\ge 0.7$ probability of being spam, for example. This is how we can trade off false positives and false negatives.

The precision-recall curve shows this trade off for each possible cutoff probability. In the cell below, [plot a precision-recall curve](http://scikit-learn.org/stable/auto_examples/model_selection/plot_precision_recall.html#plot-the-precision-recall-curve) for your final classifier (the one you use to make predictions for Kaggle).

```python
def get_feature_matrix(df):
    subject_upper = feature_subject_upper(df)
    subject_length = feature_context_length(df)
    subject_words = feature_subject_words(df)
    subject_exclaimation = feature_subject_exclaimation(df)
    subject_re = feature_re(df)
    context_length = feature_context_length(df)
    context_words = feature_context_words(df)
    html_tag = feature_html_tag(df)
    has_day = feature_has_day(df, day)
    has_month = feature_has_month(df, month)

    feature = pd.concat([subject_upper, subject_length, subject_words, subject_exclaimation, \
                         subject_re, context_length, context_words, html_tag, has_day, \
                         has_month], axis=1).values
    some_words = ["html", "head", "body", "please", "offer", "money", "sale"]
    Phi_train = words_in_texts(some_words, df["email"])
    return np.concatenate([feature, Phi_train], axis=1)
```

```python
train, val = train_test_split(original_training_data, test_size=0.1, random_state=42)

X_train = get_feature_matrix(train)
X_val = get_feature_matrix(val)
Y_train = train["spam"].values
Y_val = val["spam"].values
```

```python
lgmodel = LogisticRegression()
lgmodel.fit(X_train, Y_train)
print (lgmodel.score(X_train, Y_train))
print (lgmodel.score(X_val, Y_val))
```

<!-- grade_id: q9-roc-curve, points: 3, grade: true, solution: true -->
```python
from sklearn.metrics import precision_recall_curve, average_precision_score
# Note that you'll want to use the .predict_proba(...) method for your classifier
# instead of .predict(...) so you get probabilities, not classes

# YOUR CODE HERE
# raise NotImplementedError()
y_score = lgmodel.decision_function(X_train)
average_precision = average_precision_score(Y_train, y_score)
precision, recall, _ = precision_recall_curve(Y_train, y_score)

plt.step(recall, precision, color='b', alpha=0.2, where='post')
plt.xlabel('Recall')
plt.ylabel('Precision')
plt.ylim(0, 1)
plt.xlim(0, 1)
plt.title('Precision-Recall curve: AP={0:0.2f}'.format(average_precision))
```

# Question 10: Submitting to Kaggle

The following code will write your predictions on the test dataset to a CSV, which you can submit to Kaggle. You may need to modify it to suit your needs.

Save your predictions in a 1-dimensional array called `test_predictions`. *Even if you are not submitting to Kaggle, please make sure you've saved your predictions to `test_predictions` as this is how your grade for this part will be determined.*

Remember that if you've performed transformations or featurization on the training data, you must also perform the same transformations on the test data in order to make predictions. For example, if you've created features for the words "drug" and "money" on the training data, you must also extract the same features in order to use scikit-learn's `.predict(...)` method.

You should submit your CSV files to https://www.kaggle.com/t/d9a7013e7fd048c291ff7efe6e1ac25e

<!-- grade_id: q10-answer, points: 0, grade: false, solution: true -->
```python
# CHANGE ME (Currently making random predictions)
X_test = get_feature_matrix(test)
test_predictions = lgmodel.predict(X_test)

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: q10-tests, points: 1, grade: true, solution: false -->
```python
# must be ndarray of predictions
assert isinstance(test_predictions, np.ndarray) 

# must be binary labels (0 or 1) and not probabilities
assert np.all((test_predictions == 0) | (test_predictions == 1))

# must be the right number of predictions
assert test_predictions.shape == (1000, )
```

<!-- grade_id: q10-hidden-tests, points: 4, grade: true, solution: false -->
```python
# Please do not modify this cell
```

The following saves a file to submit to Kaggle.

```python
from datetime import datetime

# Assuming that your predictions on the test set are stored in a 1-dimensional array called
# test_predictions. Feel free to modify this cell as long you create a CSV in the right format.

# must be ndarray of predictions
assert isinstance(test_predictions, np.ndarray) 

# must be binary labels (0 or 1) and not probabilities
assert np.all((test_predictions == 0) | (test_predictions == 1))

# must be the right number of predictions
assert test_predictions.shape == (1000, )

# Construct and save the submission:
submission_df = pd.DataFrame({
    "Id": test['id'], 
    "Class": test_predictions,
}, columns=['Id', 'Class'])
timestamp = datetime.isoformat(datetime.now()).split(".")[0]
submission_df.to_csv("submission_{}.csv".format(timestamp), index=False)

print('Created a CSV file: {}.'.format("submission_{}.csv".format(timestamp)))
print('You may now upload this CSV file to Kaggle for scoring.')
```

## Submission

You're done!

Before submitting this assignment, ensure to:

1. Restart the Kernel (in the menubar, select Kernel->Restart & Run All)
2. Validate the notebook by clicking the "Validate" button

Finally, make sure to **submit** the assignment via the Assignments tab in Datahub